Aberration correction in long GRIN lens-based microendoscopes for extended field-of-view two-photon imaging in deep brain regions

Author response:

The following is the authors’ response to the previous reviews

Reviewer #1:

(1) As discussed in review and nicely simulated by the authors, the large figure error indicated by profilometry (~10 um in some cases on average) is inconsistent with the optical performance improvements observed, suggesting that those measurements are inaccurate.

I see no reason to include these inaccurate measurements.

We agree with the Referee and removed the indicated figure (old Supplementary Fig. 4) and data.

Reviewer #3:

(1) It would be interesting to comment on how the addition of a coverslip changes the performance of the uncorrected microendoscope compared to the use of bare grin lenses.

We modified the discussion section (page 18) and added a new reference (#36) to include the request of the Referee.

(2) In Figure 6C-H, the authors can indeed show data corresponding to all detected cells, but I still think that the statistics should be calculated using the same effective FOV.

We modified Figure 6 legend to include the request of the Referee.

(3) Authors could present the images in Figures 4-6 as in the original version, with a scale bar in the centre of the FOV that is different for the two types of objectives (corrected vs uncorrected). They could add a short justification for this choice, and perhaps present the other version for Figure 4 in a supplementary information sheet (with similar scale bars at the centre of the FOV for both types of objectives). It would allow readers to appreciate that the FOV still appears significantly enlarged with this other presentation.

As requested by the Referee, we modified the text in the Result section (page 11) and added the additional version of Figure 4 as Figure 4-figure supplement 1.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Two-photon (2P) fluorescence imaging through gradient index (GRIN) lens-based endoscopes is fundamental to investigate the functional properties of neural populations in deep brain circuits. However, GRIN lenses have intrinsic optical aberrations, which severely degrade their imaging performance. GRIN aberrations decrease the signal-to-noise ratio (SNR) and spatial resolution of fluorescence signals, especially in lateral portions of the field-of-view (FOV), leading to restricted FOV and smaller number of recorded neurons. This is especially relevant for GRIN lenses of several millimeters in length, which are needed to reach the deeper regions of the rodent brain. We have previously demonstrated a novel method to enlarge the FOV and improve the spatial resolution of two-photon microendoscopes based on GRIN lenses of length < 4.1 mm (Antonini et al. eLife 2020). However, previously developed microendoscopes were too short to reach the most ventral regions of the mouse brain. In this study, we combined optical simulations with fabrication of aspherical polymer microlenses through three-dimensional (3D) microprinting to correct for optical aberrations in long (length > 6 mm) GRIN lens-based microendoscopes (diameter, 500 µm). Long corrected microendoscopes had improved spatial resolution, enabling imaging in significantly enlarged FOVs. Moreover, using synthetic calcium data we showed that aberration correction enabled detection of cells with higher SNR of fluorescent signals and decreased cross-contamination between neurons. Finally, we applied long corrected microendoscopes to perform large-scale and high precision recordings of calcium signals in populations of neurons in the olfactory cortex, a brain region laying approximately 5 mm from the brain surface, of awake head-tethered mice. Long corrected microendoscopes are powerful new tools enabling population imaging with unprecedented large FOV and high spatial resolution in the most ventral regions of the mouse brain.

Revised version: Two-photon (2P) fluorescence imaging through gradient index (GRIN) lens-based endoscopes is fundamental to investigate the functional properties of neural populations in deep brain circuits. However, GRIN lenses have intrinsic optical aberrations, which severely degrade their imaging performance. GRIN aberrations decrease the signal-to-noise ratio (SNR) and spatial resolution of fluorescence signals, especially in lateral portions of the field-of-view (FOV), leading to restricted FOV and smaller number of recorded neurons. This is especially relevant for GRIN lenses of several millimeters in length, which are needed to reach the deeper regions of the rodent brain. We have previously demonstrated a novel method to enlarge the FOV and improve the spatial resolution of 2P microendoscopes based on GRIN lenses of length < 4.1 mm (Antonini et al. eLife 2020). However, previously developed microendoscopes were too short to reach the most ventral regions of the mouse brain. In this study, we combined optical simulations with fabrication of aspherical polymer microlenses through three-dimensional (3D) microprinting to correct for optical aberrations in long (length > 6 mm) GRIN lens-based microendoscopes (diameter, 500 µm). Long corrected microendoscopes had improved spatial resolution, enabling imaging in significantly enlarged FOVs. Moreover, using synthetic calcium data we showed that aberration correction enabled detection of cells with higher SNR of fluorescent signals and decreased cross-contamination between neurons. Finally, we applied long corrected microendoscopes to perform large-scale and high precision recordings of calcium signals in populations of neurons in the olfactory cortex, a brain region laying approximately 5 mm from the brain surface, of awake head-tethered mice. Long corrected microendoscopes are powerful new tools enabling population imaging with unprecedented large FOV and high spatial resolution in the most ventral regions of the mouse brain.

Introduction

High resolution 2P fluorescence imaging of the awake brain is a fundamental tool to investigate the relationship between the structure and the function of brain circuits 1. Compared to electrophysiological techniques, functional imaging in combination with genetically encoded indicators allows monitoring the activity of genetically targeted cell types, access to subcellular compartments, and tracking the dynamics of many biochemical signals in the brain 2. However, a critical limitation of multiphoton microscopy lies in its limited (< 1 mm) penetration depth in scattering biological media 3.

Removed in the revised version: Deep ventral regions of the brain, such as the olfactory cortex, the hypothalamus, and the amygdala play fundamental roles in controlling the processing of sensory information 4, circadian rhythms 5, hormone release 5, decision making 6, and adaptation of instinctive and motivational behavior 6.

Earlier version: Since these ventral areas lay deeper than 4 mm down the brain surface in rodents and larger mammals (e.g., rats, cats, and marmosets), imaging these regions typically requires the use of light guides, which relay the focal plane of the microscope objective down to the target brain area. Since inserting light guides in the brain damages the tissue above the target area 7, reducing the cross section of the probe is desired when imaging these deep brain regions. One popular solution for deep 2P imaging in awake animals is using GRIN lenses 8–11, thin cylindrical rod lenses characterized by refractive index which varies with the radial distance 12. For example, in the mouse orbitofrontal cortex 2P microscopy through GRIN endoscopes showed that orbitofrontal neurons respond to either caloric rewards or social stimuli, indicating the presence of functionally distinct feeding and social neuronal subnetworks in this deep cortical region 13. Similarly, in the lateral hypothalamus combining GRIN endoscopes with 2P imaging revealed neurons encoding thermal punishment and reward and this subset of neurons was distinct from the ensemble of neurons encoding caloric reward 14. Other ventral brain regions, such as the olfactory cortex, can be imaged with an invasive preparation, which does not require a GRIN lens but which is incompatible with imaging in awake behaving animals 15, 16. In the olfactory cortex, GRIN lens-based 2P endoscopy becomes necessary in awake animals 17, where head tilting is not compatible with animal performance in a behavioral task.

Revised version: Deep ventral regions of the brain, such as the olfactory cortex, the hypothalamus, and the amygdala play fundamental roles in controlling the processing of sensory information 4, circadian rhythms 5, hormone release 5, decision making 6, and adaptation of instinctive and motivational behavior 6. Since these ventral areas lay deeper than 4 mm down the brain surface in rodents and larger mammals (e.g., rats, cats, and marmosets), imaging these regions typically requires the use of light guides, which relay the focal plane of the microscope objective down to the target brain area. Since inserting light guides in the brain damages the tissue above the target area 7, reducing the cross section of the probe is desired when imaging these deep brain regions. One popular solution for deep 2P imaging in awake animals is using GRIN lenses 8–11, thin cylindrical rod lenses characterized by refractive index which varies with the radial distance 12. For example, in the mouse orbitofrontal cortex 2P microscopy through GRIN endoscopes showed that orbitofrontal neurons respond to either caloric rewards or social stimuli, indicating the presence of functionally distinct feeding and social neuronal subnetworks in this deep cortical region 13. Similarly, in the lateral hypothalamus combining GRIN endoscopes with 2P imaging revealed neurons encoding thermal punishment and reward and this subset of neurons was distinct from the ensemble of neurons encoding caloric reward 14. Other ventral brain regions, such as the olfactory cortex, can be imaged with an invasive preparation, which does not require a GRIN lens, but which is incompatible with imaging in awake behaving animals 15, 16. In the olfactory cortex, GRIN lens-based 2P endoscopy becomes necessary in awake animals 17, where head tilting is not compatible with animal performance in a behavioral task.

One limitation of GRIN lens-based endoscopy is that GRIN lenses suffer from intrinsic optical aberrations 18, 19. GRIN aberrations can be both on-axis and off-axis and severely degrade the quality of 2P imaging. GRIN aberrations distort the point spread function (PSF) and enlarge the excitation volume of the focal spot, leading to reduced SNR of fluorescence signals and decreased spatial resolution 7, 20, 21. This effect is not uniform across the FOV and is particularly relevant in lateral portions of the FOV, generating uneven spatial resolution and inhomogeneous amplitude of fluorescence signals across the FOV 7, 20, 21. As a consequence of these limitations, previous studies using 2P imaging through GRIN lenses in ventral regions of the mouse brain were limited to imaging a small number of neurons with uneven spatial resolution across the FOV and low SNR in lateral portion of the FOV 13, 14, 22. Developing GRIN-based endoscopic probes with improved optical properties and more homogeneous spatial resolution, which are long enough to enable imaging in the most ventral regions of the rodent brain while maintaining a small (≤ 500 µm) cross section, is thus a compelling need.

Earlier version: Various adaptive optics methods were developed to correct optical aberrations in GRIN lenses. For example, initial efforts utilized optical compensation by low-order electrostatic membrane mirror 18. Alternatively, adaptive optics through pupil segmentation was used to efficiently correct GRIN aberrations, increasing the intensity of recorded fluorescence signals and enlarging the imaging FOV 20, 21. More recently, geometric transformation adaptive optics has been developed and applied for the correction of aberration in GRIN lenses 23. However, adaptive optics approaches require substantial modification of the optical path of the microscope 18, 20, 21, 23, needs specifically designed software controls 20, 21, 23, and may limit the temporal resolution of imaging 20, 21. Alternatively, mm-size plano-convex lenses have been combined with GRIN lenses to correct on-axis aberrations and increase the Numerical Aperture (NA) of the optical system 11. However, this method is difficult to apply to GRIN lenses of cross section smaller than 1 mm 24, because of the difficulty in manufacturing high-precision optics with small radial dimension and complex profile. Recently, a novel method based on 3D microprinting of polymer optics was developed to correct for GRIN aberrations by placing specifically designed aspherical corrective lenses at the back end of the GRIN lens 7. This approach is attractive because it is built-in on the GRIN lens and corrected microendoscopes are ready-to-use, requiring no change in the optical set-up. However, previous work demonstrated the feasibility of this method only for GRIN lenses of length < 4.1 mm 7, which are too short to reach the most ventral regions of the mouse brain. The applicability of this technology to longer GRIN lenses, which are affected by stronger optical aberrations 19, remained to be proven.

Revised version: Various adaptive optics methods were developed to correct optical aberrations in GRIN lenses. For example, initial efforts utilized optical compensation by low-order electrostatic membrane mirror 18. Alternatively, adaptive optics through pupil segmentation was used to efficiently correct GRIN aberrations, increasing the intensity of recorded fluorescence signals and enlarging the imaging FOV 20, 21 More recently, geometric transformation adaptive optics has been developed and applied for the correction of aberrations in GRIN lenses 23. However, adaptive optics approaches require substantial modification of the optical path of the microscope 18, 20, 21, 23, needs specifically designed software controls 20, 21, 23, and may limit the temporal resolution of imaging 20, 21. Alternatively, mm-size plano-convex lenses have been combined with GRIN lenses to correct on-axis aberrations and increase the Numerical Aperture (NA) of the optical system 11. However, this method is difficult to apply to GRIN lenses of cross section smaller than 1 mm 24, because of the difficulty in manufacturing high-precision optics with small radial dimension and complex profile. Recently, a novel method based on 3D microprinting of polymer optics was developed to correct for GRIN aberrations by placing specifically designed aspherical corrective lenses at the back end of the GRIN lens 7. This approach is attractive because it is built-in on the GRIN lens and corrected microendoscopes are ready-to-use, requiring no change in the optical set-up. However, previous work demonstrated the feasibility of this method only for GRIN lenses of length < 4.1 mm 7, which are too short to reach the most ventral regions of the mouse brain. The applicability of this technology to longer GRIN lenses, which are affected by stronger optical aberrations 19, remained to be proven.

In this study, we designed, developed and validated correction of optical aberrations in GRIN lenses of length 6.4 mm and 8.8 mm and diameter 500 µm using a combination of ray-trace simulation, 3D microprinting by 2P lithography, simulated calcium data, and experimental characterization. Validation of the new corrected microendoscopes is provided with proof-of-principle experiments in a ventral brain region, the olfactory cortex, of awake head-fixed mice.

Results

Design of aberration corrected microendoscopic probes based on long GRIN rods

Earlier version: We selected two GRIN singlet rod lenses (here on GRIN rods) suitable for multiphoton fluorescence imaging in ventral regions of the mouse brain. The two selected GRIN rods had NA = 0.5 on both sides, 500 µm diameter, 1.5 and 2.0 pitch, and 6.4 mm and 8.8 mm length, respectively. For both GRIN rods, we modeled an optical assembly composed of the GRIN rod, a thin layer of glass (representing a 100 µm thick glass coverslip) attached to the back end of the GRIN rod, and an aspheric corrective lens (here on called corrective lens) aligned to the rod and attached to the opposite side of the glass coverslip with respect to the GRIN rod (Fig. 1A, B). We designed the aspheric surface of the corrective lens using ray-trace simulations at the wavelength (;\exc) of 920 nm, the excitation wavelength commonly used for 2P imaging of the popular genetically encoded calcium indicator GCaMP. The design of the corrective lens was an iterative process in which the parameters that define the corrective lens surface were automatically varied to obtain the highest and most homogeneous spatial resolution possible over the longest possible radial distance from the optical axis (see Supplementary Table 1 and Materials and Methods for details on simulation parameters). More specifically, we simulated the corrective lens profile such that the Strehl ratio (defined as the peak intensity of the focal spot of an aberrated optical system normalized to the peak intensity of the focal spot of the ideal optical system 25) of the corrected microendoscope remained above the threshold of 0.8 for the longest possible radial distance from the optical axis. The values of Strehl ratio equal to 0.8 corresponded to the lower bound of the diffraction-limited condition set by the Maréchal criterion 7, 25. Our simulations showed that the profiles of the two corrective lenses displayed in Fig. 1 A, B led to an increase in the radius of the diffraction-limited FOV compared to the case in which the corrective lens was not positioned in the back end of the GRIN rods (uncorrected microendoscopes). The profile of the corrective optical element was specific to the type of GRIN rod considered. The simulated increase in the radius of the diffraction-limited FOV was 3.50 times and 3.35 times for the 6.4 mm-long and 8.8 mm-long probe, respectively (Fig. 1C, E). Moreover, we investigated the effect of changing wavelength on the Strehl ratio. We found that the Strehl ratio remained > 0.8 within at least ± 10 nm from λ = 920 nm (Supplementary Figure 1), which covered the limited bandwidth of our femtosecond laser. All simulations were performed maximizing aberration correction in the focal plane of the microendoscopes. However, we explored the effect of aberration correction outside the focal plane. In corrected endoscopes, we found that for off-axis rays (radial distance from the optical axis > 40 μm) the Strehl ratio was > 0.8 (Maréchal criterion) in a larger volume compared to uncorrected endoscopes (new Supplementary Figure 2), demonstrating that the aberration correction method developed in this study extends for short distances beyond the focal plane. For example, at a radial distance of ∼ 90 μm from the optical axis, the axial range in which the Strehl ratio was > 0.8 in corrected endoscopes was 28 μm and 19 μm for the 6.4 mm-and the 8.8 mm-long microendoscope, respectively. Decentering the optical axis of the corrective lens from the optical axis of the GRIN rod rapidly disrupted the optical properties of the corrected microendoscopes (Supplementary Figure 3). To visualize the expected improvement of the spatial resolution in corrected microendoscopes, we simulated the PSF of uncorrected and corrected microendoscopes at different radial distances from the optical axis. For both the 6.4 mm-long and 8.8 mm-long GRIN rods, we found that the PSF of the uncorrected probe was strongly aberrated and highly irregular at ∼ 90 µm away from the center of the FOV (Fig. 1D, F). In contrast, in the corrected microendoscopes the PSF had regular shape and, at > 90 µm away from the optical axis, it was similar to the one measured at the center of the FOV (Fig. 1D, F, see also Supplementary Table 2).

Revised version: We selected two GRIN singlet rod lenses (here on GRIN rods) suitable for multiphoton fluorescence imaging in ventral regions of the mouse brain. The two selected GRIN rods had NA = 0.5 on both sides, 500 µm diameter, 1.5 and 2.0 pitch, and 6.4 mm and 8.8 mm length, respectively. For both GRIN rods, we modeled an optical assembly composed of the GRIN rod, a thin layer of glass (representing a 100 µm thick glass coverslip) attached to the back end of the GRIN rod, and an aspheric corrective lens (here on called corrective lens) aligned to the rod and attached to the opposite side of the glass coverslip with respect to the GRIN rod (Figure 1A, B). We designed the aspheric surface of the corrective lens using ray-trace simulations at the wavelength (λexc) of 920 nm, the excitation wavelength commonly used for 2P imaging of the popular genetically encoded calcium indicator GCaMP. The design of the corrective lens was an iterative process in which the parameters that define the corrective lens surface were automatically varied to obtain the highest and most homogeneous spatial resolution possible over the longest possible radial distance from the optical axis (see Supplementary Table 1 and Materials and Methods for details on simulation parameters). More specifically, we simulated the corrective lens profile such that the Strehl ratio (defined as the peak intensity of the focal spot of an aberrated optical system normalized to the peak intensity of the focal spot of the ideal optical system 25) of the corrected microendoscope remained above the threshold of 0.8 for the longest possible radial distance from the optical axis. The values of Strehl ratio equal to 0.8 corresponded to the lower bound of the diffraction-limited condition set by the Maréchal criterion 7, 25. Our simulations showed that the profiles of the two corrective lenses displayed in Figure 1A, B led to an increase in the radius of the diffraction-limited FOV compared to the case in which the corrective lens was not positioned in the back end of the GRIN rods (uncorrected microendoscopes). The profile of the corrective optical element was specific to the type of GRIN rod considered. The simulated increase in the radius of the diffraction-limited FOV was 3.50 times and 3.35 times for the 6.4 mm-long and 8.8 mm-long probe, respectively (Figure 1C, E). Moreover, we investigated the effect of changing wavelength on the Strehl ratio. We found that the Strehl ratio remained > 0.8 within at least ± 10 nm from λexc = 920 nm (Figure 1-figure supplement 1), which covered the limited bandwidth of our femtosecond laser. All simulations were performed maximizing aberration correction in the simulated focal plane of the microendoscopes (imaging plane in Figure 1A, B). However, we explored the effect of aberration correction outside the simulated focal plane. In corrected microendoscopes, we found that for off-axis rays (radial distance from the optical axis > 40 μm) the Strehl ratio was > 0.8 (Maréchal criterion) in a larger volume compared to uncorrected microendoscopes (Figure 1-figure supplement 2), demonstrating that the aberration correction method developed in this study extends for short distances beyond the simulated focal plane. For example, at a radial distance of ∼ 90 μm from the optical axis, the axial range in which the Strehl ratio was > 0.8 in corrected microendoscopes was 28 μm and 19 μm for the 6.4 mm-long and the 8.8 mm-long microendoscope, respectively. Decentering the optical axis of the corrective lens from the optical axis of the GRIN rod rapidly disrupted the optical properties of the corrected microendoscopes (Figure 1-figure supplement 3). To visualize the expected improvement of the spatial resolution in corrected microendoscopes, we simulated the PSF of uncorrected and corrected microendoscopes at different radial distances from the optical axis. For both the 6.4 mm-long and 8.8 mm-long GRIN rods, we found that the PSF of the uncorrected probe was strongly aberrated and highly irregular at ∼ 90 µm away from the center of the FOV (Figure 1D, F). In contrast, in the corrected microendoscopes the PSF had regular shape and, at > 90 µm away from the optical axis, it was similar to the one measured at the center of the FOV (Figure 1D, F, see also Supplementary Table 2).

Optical simulations of long corrected microendoscopes.

A) Ray-trace simulation for the microendoscope based on the 6.4 mm-long GRIN rod. Left: rays of 920 nm light are relayed from the objective focal plane to the imaging plane. Labels indicate geometrical parameters of the microendoscope components. Right: profile of the corrective aspherical lens that maximizes the microendoscope FOV in the optical simulation shown on the left. B) Same as (A) for the microendoscope based on the 8.8 mm-long GRIN rod. C,D) Optical performance of simulated microendoscope based on the 6.4 mm-long GRIN rod. C) Strehl ratio as a function of the field radial distance (zero indicates the optical axis) computed on the focal plane in the object space (after the GRIN rod) for the corrected (blue) and the uncorrected (red) microendoscope. The black horizontal dashed line indicates the diffraction-limited threshold according to the Maréchal criterion 25. The black vertical dashed lines mark the abscissa values of the intersections between the curves and the diffraction-limited threshold. D) Lateral (x,y) and axial (x,z and y,z) intensity profiles of simulated PSFs on-axis (distance from the center of the FOV d = 0 µm) and off-axis (at the indicated distance d) for the uncorrected (left) and the corrected microendoscope (right). Horizontal scale bars: 1 µm; vertical scale bars: 10 µm. E,F) Same as (C,D) for the microendoscope based on the 8.8 mm-long GRIN rod.

Microfabrication of corrective lenses and assembly of long corrected microendoscopes

Earlier version: Due to the small size of the simulated corrective lenses (diameter: 500 µm; height: few tens of µm), we first used 3D microprinting based on 2P lithography to fabricate aspherical corrective lens prototypes (Fig. 2A, B) 26, 27. We then replicated the prototypes using a molding procedure 28 (Fig. 2C). The molding strategy enabled us to obtain large numbers of aspherical lens replica directly manufactured onto a round glass coverslip in a much faster and cheaper way compared to 3D microprinting. We checked that the surface of the corrective lenses was in good agreement with the simulated profile (Supplementary Figure 4). Next, we assembled microendoscopes following the scheme in Fig. 2D, using a plain coverslip for the uncorrected case, while we used a coverslip with a corrective lens replica for the corrected case (this applies to all the experiments described in the manuscript). In the assembly of the microendoscope, we attached the round coverslip to the annular support structure shown in (Fig. 2D, called hereon ring), which had two functions: i) to facilitate holding the probe during stereotactic surgery implantation; ii) to provide stability to the microendoscope once implanted on the animal skull. We used either metallic rings (external diameter: 7 mm) or plastic rings (external diameter: 4.5 mm). Plastic rings were specifically designed and 3D printed to be compatible with commercial holders (ProView Implant Kit, Inscopix Inc. Mountain View, CA) to ease stereotaxic implantation (see Materials and Methods).

Revised version: Due to the small size of the simulated corrective lenses (diameter: 500 µm; height: few tens of µm), we first used 3D microprinting based on 2P lithography to fabricate aspherical corrective lens prototypes (Figure 2A, B) 26, 27. We then replicated the prototypes using a molding procedure 28 (Figure 2C). The molding strategy enabled us to obtain large numbers of aspherical lens replica directly manufactured onto round glass coverslips in a much faster and cheaper way compared to 3D microprinting. Next, we assembled microendoscopes following the scheme in Figure 2D, using a plain coverslip for the uncorrected case, while we used a coverslip with a corrective lens replica for the corrected case (this applies to all the experiments described in the manuscript). In the assembly of the microendoscope, we attached the round coverslip to the annular support structure shown in (Figure 2D, called hereon ring), which had two functions: i) to facilitate holding the probe during stereotactic surgery implantation; ii) to provide stability to the microendoscope once implanted on the animal skull. We used either metallic rings (external diameter: 7 mm) or plastic rings (external diameter: 4.5 mm). Plastic rings were specifically designed and 3D printed to be compatible with commercial holders (ProView Implant Kit, Inscopix Inc. Mountain View, CA) to ease stereotaxic implantation (see Materials and Methods). The alignment and assembly of all the microendoscope components was performed with a custom optomechanical set-up similar to that described in 7.

Fabrication of corrective lenses using 3D microprinting and assembly of long microendoscopes.

Earlier version: A) Scanning electron microscopy image of the 3D microprinted replica of the corrective lens for the 6.4 mm-long corrected microendoscope. Scale bar: 100 µm. B) Same as (A) for the 8.8 mm-long microendoscope. C) Molding procedure for the generation of corrective lens replica. Freshly prepared PDMS is casted onto a corrective aspherical lens printed with 2P lithography (1). After 48 hours, the solidified PDMS provides a negative mold for the generation of lens replica (2). A small drop of optical UV-curable glue is deposited onto the mold (3). The mold filled by glue is covered with a glass coverslip, which is gently pressed against the mold (4). The optical glue is polymerized with UV light (5). The coverslip with the attached polymeric lens replica is detached from the negative mold (6). Object dimensions are not to scale. D) Exploded view of the corrected microendoscope assembly.

Revised version: A) Scanning electron microscopy image of the 3D microprinted replica of the corrective lens for the 6.4 mm-long corrected microendoscope. Scale bar: 100 µm. B) Same as (A) for the 8.8 mm-long microendoscope. C) Molding procedure for the generation of corrective lens replica. Freshly prepared PDMS is casted onto a corrective aspherical lens printed with 2P lithography (1). After 48 hours, the solidified PDMS provides a negative mold for the generation of lens replica (2). A small drop of optical UV-curable glue is deposited onto the mold (3). The mold filled by optical glue is covered with a glass coverslip, which is gently pressed against the mold (4). The optical glue is polymerized with UV light (5). The coverslip with the attached polymeric lens replica is detached from the negative mold (6). Object dimensions are not to scale. D) Exploded view of the corrected microendoscope assembly.

Removed in the revised version: The alignment and assembly of all the microendoscope components was performed with a custom optomechanical set-up similar to that described in 7.

Improved optical performance of long corrected microendoscopes

Earlier version: We tested the optical performance of the new corrected microendoscopes in 2P laser scanning microscopy (2PLSM) using a pulsed laser tuned at 920 nm. Microendoscopes were coupled to the 2P microscope using a previously described customized mount 7. To visualize the effect of the corrective lens on the axial extension of the excitation volume, we imaged subresolved homogenous fluorescent layers (layer thickness: ∼300 nm) using uncorrected and corrected microendoscopes. Figure 3A-D shows the x,z intensity projections of z-stacks acquired through the subresolved fluorescent layers under both conditions (i.e., uncorrected and corrected microendoscopes). Strong off-axis aberrations of the GRIN lens generated a significant loss of fluorescence intensity in the marginal portions of the FOV in uncorrected microendoscopes (both the 6.4 mm-long, Fig. 3A, B, and the 8.8 mm-long, Fig 3C, D). In contrast, fluorescence signal was collected over longer distances from the center of the FOV in corrected microendoscopes (Fig. 3A-D). Moreover, for GRIN rods of both lengths the full width at half maximum (FWHM) of the film profile obtained with uncorrected microendoscopes increased with distance from the optical axis and became rapidly larger than 20 µm. In contrast, the film profile obtained with corrected microendoscopes showed more homogenous FWHM values across an extended FOV (Fig. 3A-D).

Revised version: We tested the optical performance of the new corrected microendoscopes in 2P laser scanning microscopy (2PLSM) using a pulsed laser tuned at 920 nm. Microendoscopes were coupled to the 2P microscope using a previously described customized mount 7. To visualize the effect of the corrective lens on the axial extension of the excitation volume, we imaged subresolved homogenous fluorescent layers (layer thickness: ∼300 nm) using uncorrected and corrected microendoscopes. Figure 3A-D shows the x,z intensity projections of z-stacks acquired through the subresolved fluorescent layers under both conditions (i.e., uncorrected and corrected microendoscopes). Strong off-axis aberrations of the GRIN lens generated a significant loss of fluorescence intensity in the marginal portions of the FOV in uncorrected microendoscopes (for both the 6.4 mm-long, Figure 3A, B, and the 8.8 mm-long, Figure 3C, D). In contrast, fluorescence signal was collected over longer distances from the center of the FOV in corrected microendoscopes (Figure 3A-D). Moreover, for GRIN rods of both lengths the full width at half maximum (FWHM) of the film profile obtained with uncorrected microendoscopes increased with distance from the optical axis and became rapidly larger than 20 µm. In contrast, the film profile obtained with corrected microendoscopes showed more homogenous FWHM values across an extended FOV (Figure 3A-D).

Optical characterization of long corrected microendoscopes shows improved spatial resolution over an enlarged FOV.

Earlier version: A) Representative x (horizontal), z (vertical) projection of a z-stack of a subresolved fluorescent layer acquired with the uncorrected (top) or corrected (bottom) microendoscope based on the 6.4 mm-long GRIN rod. λexc = 920 nm; scale bars: 50 pixels. B) Thickness (mean values ± standard error of the mean (s.e.m.)) of the layer as a function of the distance from the center of the FOV for uncorrected (red, n = 4) or corrected (blue, n = 4) microendoscopes. The thickness of the film is measured as the FWHM of the Gaussian fit of the fluorescence intensity along segments orthogonal to the tangential line to the section of the film and located at different distances from the center of the FOV (see yellow labels on the bottom image). C,D) Same as (A,B) for the microendoscope based on the 8.8 mm-long GRIN rod. E) The distortion of the FOV in uncorrected and corrected microendoscopes is evaluated using a calibration ruler. The magnification factor is defined as the ratio between the nominal and the real pixel size of the image and shown as a function of the radial distance for uncorrected (red, n = 3) or corrected (blue, n = 3) microendoscopes. Data are shown as mean values ± s.e.m.. Fitting curves are quartic functions f(x) = ax4 + bx2 + c (see also Supplementary Table 3 for details). F) Same as (E) for the microendoscope based on the 8.8 mm-long GRIN rod. G-J) The spatial resolution of microendoscopes was measured acquiring z-stacks of subresolved fluorescent beads (bead diameter: 100 nm) located at different radial distances using 2PLSM (λexc = 920 nm). G) Representative x,y and x,z projections of a fluorescent bead located at a radial distance of 75 μm, imaged through an uncorrected (left) or a corrected (right) 6.4 mm-long microendoscope. Horizontal scale bars, 2 μm; vertical scale bars, 5 μm. H) Same as G) for the microendoscope based on the 8.8 mm-long GRIN rod. I) Axial (left) and lateral (right) resolution (i.e. average size of the x,z and x,y projections of imaged beads, respectively) as a function of the radial distance from the center of the FOV for uncorrected (red) and corrected (blue) probes. Each data point represents the mean value ± s.e.m. of n = 4-24 beads imaged using at least m = 3 different 6.4 mm-long microendoscopes. Fitting curves are quartic functions f(x) = ax4 + bx2 + c (see Supplementary Table 4for details). The horizontal black dash-dotted line indicates the axial resolution threshold of 10 µm. The black triangles indicate the intersections between the threshold and the curves fitting the data and mark the estimated radius of the effective FOV of the probes. J) Same as (I) for the microendoscopes based on the 8.8 mm-long GRIN rod.

Revised version: A) Representative x (horizontal), z (vertical) projection of a z-stack of a subresolved fluorescent layer acquired with the uncorrected (top) or corrected (bottom) microendoscope based on the 6.4 mm-long GRIN rod. λexc = 920 nm; scale bars: 50 pixels. B) Thickness (mean values ± standard error of the mean (s.e.m.)) of the layer as a function of the distance from the center of the FOV for uncorrected (red, n = 4) or corrected (blue, n = 4) microendoscopes. The thickness of the film is measured as the FWHM of the Gaussian fit of the fluorescence intensity along segments orthogonal to the tangential line to the section of the film and located at different distances from the center of the FOV (see yellow labels on the bottom image). C,D) Same as (A,B) for the microendoscope based on the 8.8 mm-long GRIN rod. E) The distortion of the FOV in uncorrected and corrected microendoscopes is evaluated using a calibration ruler. The magnification factor is defined as the ratio between the nominal and the real pixel size of the image and shown as a function of the radial distance for uncorrected (red, n = 3) or corrected (blue, n = 3) microendoscopes. Data are shown as mean values ± s.e.m.. Fitting curves are quartic functions f(x) = ax4 + bx2 + c (see also Supplementary Table 3 for details). F) Same as (E) for the microendoscope based on the 8.8 mm-long GRIN rod. G-J) The spatial resolution of microendoscopes was measured acquiring z-stacks of subresolved fluorescent beads (bead diameter: 100 nm) located at different radial distances using 2PLSM (λexc = 920 nm). G) Representative x,y and x,z projections of a fluorescent bead located at a radial distance of 75 μm, imaged through an uncorrected (left) or a corrected (right) 6.4 mm-long microendoscope. Horizontal scale bars, 2 μm; vertical scale bars, 5 μm. H) Same as G) for the microendoscope based on the 8.8 mm-long GRIN rod. I) Axial (left) and lateral (right) resolution (i.e. average size of the x,z and x,y projections of imaged beads, respectively) as a function of the radial distance from the center of the FOV for uncorrected (red) and corrected (blue) probes. Each data point represents the mean value ± s.e.m. of n = 4-24 beads imaged using at least m = 3 different 6.4 mm-long microendoscopes. Fitting curves are quartic functions f(x) = ax4 + bx2 + c (see Supplementary Table 4 for details). The horizontal black dash-dotted line indicates the axial resolution threshold of 10 µm. The black triangles indicate the intersections between the threshold and the curves fitting the data and mark the estimated radius of the effective FOV of the probes. J) Same as (I) for the microendoscopes based on the 8.8 mm-long GRIN rod.

Earlier version: To evaluate the increase in area of the sample that could be imaged with the corrected microendoscopes, we first characterized potential distortion of the FOV introduced by the microendoscopes. Using a calibration ruler, we measured the local magnification factor at different radial distances from the optical axis (Fig. 3E, F). We defined the local magnification factor as the ratio between the nominal pixel size (i.e. pixel size that applies to undistorted images collected with the microscope objective alone) and the real pixel size of the image collected with the same microendoscope objective coupled to the microendoscope (see Materials and Methods). Pixel size and distance displayed in the following panels and graphs are calibrated for FOV distortion and indicate real dimensions in the samples (see also Materials and Methods and Supplementary Table 3). We then measured the spatial resolution of microendoscopes by imaging subresolved fluorescent beads (bead diameter: 100 nm) located at different radial distances from the optical axis (Fig. 3G, J). We found that the axial dimension of the beads imaged with uncorrected microendoscopes rapidly increased with radial distance, whereas the spatial resolution of corrected microendoscopes remained more homogeneous over longer radial distances and more similar to the one measured on-axis (Fig. 3G, J and Supplementary Table 4). We set a threshold of 10 µm on the axial resolution to define the radius of the effective FOV (corresponding to the black triangles in Fig. 3I, J) in uncorrected and corrected microendoscopes. We observed a relative increase of the effective FOV radius of 2.17 and 1.53 for the 6.4 mm-and the 8.8 mm-long microendoscope, respectively (Table 1). This corresponded to an enlargement of the effective FOV area of 4.7 times and 2.3 times for the 6.4 mm-long microendoscope and the 8.8 mm-long microendoscope, respectively (Table 1). These findings were in agreement with the results of the ray-trace simulations (Figure 1) and the measurement of the subresolved fluorescence layers (Figure 3A- D).

Revised version: To evaluate the increase in area of the sample that could be imaged with the corrected microendoscopes, we first characterized potential distortion of the FOV introduced by the microendoscopes. Using a calibration ruler, we measured the local magnification factor at different radial distances from the optical axis (Figure 3E, F). We defined the local magnification factor as the ratio between the nominal pixel size (i.e. pixel size that applies to undistorted images collected with the microscope objective alone) and the real pixel size of the image collected with the same objective coupled to the microendoscope (see Materials and Methods). Pixel sizes and distances displayed in the following panels and graphs are calibrated for FOV distortion and indicate real dimensions in the samples (see also Materials and Methods and Supplementary Table 3). We then measured the spatial resolution of microendoscopes by imaging subresolved fluorescent beads (bead diameter: 100 nm) located at different radial distances from the optical axis (Figure 3G-J). We found that the axial dimension of the beads imaged with uncorrected microendoscopes rapidly increased with radial distance, whereas the spatial resolution of corrected microendoscopes remained more homogeneous over longer radial distances and more similar to the one measured on-axis (Figure 3G-J and Supplementary Table 4). We set a threshold of 10 µm on the axial resolution to define the radius of the effective FOV (corresponding to the black triangles in Figure 3I, J) in uncorrected and corrected microendoscopes. We observed a relative increase of the effective FOV radius of 2.17 and 1.53 for the 6.4 mm-long and the 8.8 mm-long microendoscope, respectively (Table 1). This corresponded to an enlargement of the effective FOV area of 4.7 times and 2.3 times for the 6.4 mm-long microendoscope and the 8.8 mm-long microendoscope, respectively (Table 1). These findings were in agreement with the results of the ray-trace simulations (Figure 1) and the measurement of the subresolved fluorescence layers (Figure 3A-D).

Earlier version: The optical characterization described above indicated that corrected microendoscopes of both lengths enabled imaging a larger FOV area compared to uncorrected microendoscopes. This result can be explained by the fact that corrected microendoscopes showed higher spatial resolution and increased probability of 2P fluorescence excitation in the marginal portions of the FOV. Corrected microendoscopes should therefore allow highly contrasted imaging of cell bodies and thin cellular processes over an extended FOV compared to uncorrected microendoscopes. Using both uncorrected and corrected microendoscopes, we confirmed these predictions by imaging the same region of a fixed brain slice expressing jGCaMP7f in neurons. We found that aberration correction enabled to clearly resolve neuronal somata and processes in the lateral portions of the FOV, whereas in the absence of the corrective lens the same structures located in the peripheral portion of the FOV appeared dim, strongly aberrated, and hardly detectable (Fig. 4).

Revised version: The optical characterization described above indicated that corrected microendoscopes of both lengths enabled imaging a larger FOV area compared to uncorrected microendoscopes. This result can be explained by the fact that corrected microendoscopes showed higher spatial resolution and increased probability of 2P fluorescence excitation in the marginal portions of the FOV. Corrected microendoscopes should therefore allow highly contrasted imaging of cell bodies and thin cellular processes over an extended FOV compared to uncorrected microendoscopes. Using both uncorrected and corrected microendoscopes, we confirmed these predictions by imaging the same region of a fixed brain slice expressing jGCaMP7f in neurons. We found that aberration correction enabled to clearly resolve neuronal somata and processes in the lateral portions of the FOV, whereas in the absence of the corrective lens the same structures located in the peripheral portion of the FOV appeared dim, strongly aberrated, and hardly detectable (Figure 4, see also Figure 4-figure supplement 1 for an alternative visualization in which the FOVs of corrected microendoscopes are rescaled to match the real pixel size of the FOVs of uncorrected microendoscopes in the center of the image).

Aberration correction in long GRIN lens-based microendoscopes enables high-resolution imaging of biological structures over enlarged FOVs.

Earlier version: A) jGCaMP7f-stained neurons in a fixed mouse brain slice were imaged using 2PLSM (λexc = 920 nm) through an uncorrected (left) and a corrected (right) microendoscope based on the 6.4 mm-long GRIN rod. Images are maximum fluorescence intensity (F) projections of a z-stack acquired with a 5 μm step size. Number of steps: 32 and 29 for uncorrected and corrected microendoscope, respectively. Scale bars: 50 μm. Left: the scale applies to the entire FOV. Right, the scale bar refers only to the center of the FOV; off-axis scale bar at any radial distance (x and y axes) is locally determined multiplying the length of the drawn scale bar on-axis by the corresponding normalized magnification factor shown in the horizontal color-coded bar placed below the image (see also Fig. 3, Supplementary Table 3, and Materials and Methods for more details). B) Same results for the microendoscope based on the 8.8 mm-long GRIN rod. Number of steps: 23 and 31 for uncorrected and corrected microendoscope, respectively.

Revised version: A) jGCaMP7f-stained neurons in a mouse fixed brain slice were imaged using 2PLSM (λexc = 920 nm) through an uncorrected (left) and a corrected (right) microendoscope based on the 6.4 mm-long GRIN rod. Images are maximum fluorescence intensity (F) projections of a z-stack acquired with a 5 μm step size. Number of steps: 32 and 29 for uncorrected and corrected microendoscope, respectively. Scale bars: 50 μm. Left: the scale applies to the entire FOV. Right: the scale bar refers only to the center of the FOV; off-axis scale bar at any radial distance (x and y axes) is locally determined multiplying the length of the drawn scale bar on-axis by the corresponding normalized magnification factor shown in the horizontal color-coded bar placed below the image (see also Figure 3, Supplementary Table 3, and Materials and Methods for more details). B) Same results for the microendoscope based on the 8.8 mm-long GRIN rod. Number of steps: 23 and 31 for uncorrected and corrected microendoscope, respectively.

More accurate sampling of simulated neuronal activity with long corrected microendoscopes

Earlier version: Using simulated calcium t-series, we next evaluated the impact of the improved optical performance of long aberration corrected microendoscopes on the ability to extract information about neuronal activity. To this aim, we compared simulated calcium traces imaged with either uncorrected or corrected microendoscopes with the ground truth neuronal activity, i.e. the neuronal activity used to generate the simulated t-series. To build synthetic calcium data, we first generated neurons with 3D distribution and anatomical properties (i.e., cell size and cell density) similar to those measured in the mouse olfactory cortex 29, 30, the ventral brain region which we will be imaging in vivo (see next section of Results), and the biophysical characteristics of one of the two indicators used in our experiments, jGCaMP8f 31 (see also Materials and Methods). We then simulated the optical sampling of 3D volume of neurons with the PSF properties of the uncorrected and corrected microendoscopes characterized above (Fig. 3). Fluorescence signals integrated over the PSF volume were then projected on a 2D matrix of pixels in time, with a frequency of 30 Hz, to obtain synthetic t-series (duration: 5 min, Fig. 5A, D). We finally detected neuronal cell bodies and extracted fluorescence traces from regions of interest (ROIs) of synthetic t-series using established methods 32 commonly applied to real data and compared between the case of uncorrected and corrected microendoscopes. We found that the use of corrected probes allowed segmenting more ROIs with high SNR compared to uncorrected microendoscopes, shifting the distribution of SNR across ROIs to higher mean SNR values. This was true for both the microendoscopes based on the 6.4 mm GRIN rod (Fig. 5B) and the microendoscope based on the 8.8 mm long GRIN rod (Fig. 5E). Moreover, for both microendoscope types, aberration correction increased the maximal radial distance at which ROIs with high (> 15) SNR could be observed (Fig. 5C, F). These findings demonstrate that correcting for optical aberrations increases the probability of recording ROIs with larger SNR of fluorescent signals.

Revised version: Using simulated calcium t-series, we next evaluated the impact of the improved optical performance of long aberration corrected microendoscopes on the ability to extract information about neuronal activity. To this aim, we compared simulated calcium traces imaged with either uncorrected or corrected microendoscopes with the ground truth neuronal activity, i.e. the neuronal activity used to generate the simulated t-series. To build synthetic calcium data, we first generated neurons with 3D distribution and anatomical properties (i.e., cell size and cell density) similar to those measured in the mouse olfactory cortex 29, 30, the ventral brain region which we will be imaging in vivo (see next section of Results), and the biophysical characteristics of one of the two indicators used in our experiments, jGCaMP8f 31 (see also Materials and Methods). We then simulated the optical sampling of 3D volume of neurons with the PSF properties of the uncorrected and corrected microendoscopes characterized above (Figure 3). Fluorescence signals integrated over the PSF volume were then projected on a 2D matrix of pixels in time, with a frequency of 30 Hz, to obtain synthetic t-series (duration: 5 min, Figure 5A, D). We finally detected neuronal cell bodies and extracted fluorescence traces from regions of interest (ROIs) of synthetic t-series using established methods 32 commonly applied to real data and compared between the case of uncorrected and corrected microendoscopes. We found that the use of corrected probes allowed segmenting more ROIs with high SNR compared to uncorrected microendoscopes, shifting the distribution of SNR across ROIs to higher mean SNR values. This was true for both the microendoscopes based on the 6.4 mm GRIN rod (Figure 5B) and the microendoscope based on the 8.8 mm long GRIN rod (Figure 5E). Moreover, for both microendoscope types, aberration correction increased the maximal radial distance at which ROIs with high (> 15) SNR could be observed (Figure 5C, F). These findings demonstrate that correcting for optical aberrations increases the probability of recording ROIs with larger SNR of fluorescent signals.

Long corrected microendoscopes sample more homogeneously simulated neuronal activity across the FOV.

Earlier version: A) Median fluorescence intensity (F) projections of representative synthetic t-series for the uncorrected (left) and the corrected (right) 6.4 mm-long microendoscope. Cell identities were detected using CITE-ON 32 in n = 13 simulated t-series for both the uncorrected and corrected microendoscopes and cellular activity traces were extracted. Green rectangular boxes mark cell identities that have peak SNR of the activity trace higher than a threshold set to peak SNR = 15. Scale bars: 50 μm. Left: the scale applies to the entire FOV. Right: the scale bar refers to the center of the FOV; off-axis scale bar at any radial distance (x and y axes) is locally determined multiplying the length of the drawn scale bar by the corresponding normalized magnification factor shown in the horizontal color-coded bar placed below the image. B) Number of detected cell identities in simulated FOV as a function of peak SNR threshold imposed on cellular activity traces. Data are mean values ± s.e.m. for both the uncorrected (red) or corrected (blue) case. Statistical significance is assessed with Mann-Whitney U test; *, p < 0.05; ***, p < 0.001. C) Maximal distance from the center of the FOV at which a cell is detected as a function of peak SNR threshold. Data are mean values ± s.e.m. for both the uncorrected (red) and corrected (blue) case. Statistical significance is assessed with Mann-Whitney U test; **, p < 0.01; ***, p < 0.001; n.s., non significant. D) Same as (A) for the 8.8 mm-long microendoscope. E, F) Same as (B,C) for n = 15 simulated t-series for both the uncorrected and corrected 8.8 mm-long microendoscope.

Revised version: A) Median fluorescence intensity (F) projections of representative synthetic t-series for the uncorrected (left) and the corrected (right) 6.4 mm-long microendoscope. Cell identities were detected using CITE-ON 32 in n = 13 simulated t-series for both the uncorrected and corrected microendoscopes and cellular activity traces were extracted. Green rectangular boxes mark cell identities that have peak SNR of the activity trace higher than a threshold set to peak SNR = 15. Scale bars: 50 μm. Left: the scale applies to the entire FOV. Right: the scale bar refers to the center of the FOV; off-axis scale bar at any radial distance (x and y axes) is locally determined multiplying the length of the drawn scale bar by the corresponding normalized magnification factor shown in the horizontal color-coded bar placed below the image. B) Number of detected cell identities in simulated FOV as a function of peak SNR threshold imposed on cellular activity traces. Data are mean values ± s.e.m. for both the uncorrected (red) or corrected (blue) case. Statistical significance is assessed with Mann-Whitney U test; *, p < 0.05; ***, p < 0.001. C) Maximal distance from the center of the FOV at which a cell is detected as a function of peak SNR threshold. Data are mean values ± s.e.m. for both the uncorrected (red) and corrected (blue) case. Statistical significance is assessed with Mann-Whitney U test; **, p < 0.01; ***, p < 0.001; n.s., not significant. D) Same as (A) for the 8.8 mm-long microendoscope. E,F) Same as (B,C) for n = 15 simulated t-series for both the uncorrected and corrected 8.8 mm-long microendoscope.

Earlier version: The higher and more homogeneous axial resolution across the FOV afforded by the corrected probes could be useful not only to increase the SNR of neural calcium signals, but also to better isolate individual neurons, decreasing signal contamination across neighboring cells. To test this hypothesis, we considered correlations in “adjacent” simulated neurons, defined as pairs of neurons for which the distance between their centroids was ≤ 25 µm (the soma of excitatory neurons in the olfactory cortex is around 10 µm 30). Compared to uncorrected probes, we found that corrected microendoscopes of both lengths led to lower fraction of pairs of adjacent cells whose activity correlated significantly more than expected (Fig. 6A, F and Supplementary Table 5), confirming reduced cross contamination among cells. The fraction of pairs of adjacent cells (out of the total number of adjacent pairs) whose activity correlated significantly more than expected increased as a function of the SNR threshold for corrected and uncorrected microendoscopes of both lengths (Fig. 6A, F). This effect was due to a larger decrease of the total number of pairs of adjacent cells as a function of the SNR threshold compared to the decrease in the number of pairs of adjacent cells whose activity was more correlated than expected (Supplementary Figure 5). Moreover, in uncorrected microendoscopes the pairwise correlation of adjacent cells showed a larger linear dependence on the radial distance of the pair from the center of the FOV compared to corrected microendoscopes (Supplementary Figure 6A, C, see also Materials and Methods). We interpreted these findings on correlated neuronal activity as due to the enlarged PSF of uncorrected microendoscopes in lateral portion of the FOV (Fig. 3), leading to increased signal contamination by sampling across neurons. These results thus suggest that corrected microendoscopes decrease the artefactual correlation of nearby neurons by diminishing the spatial extension of the excitation volume used for sampling neuronal activity (PSF, Fig. 3).

Revised version: The higher and more homogeneous axial resolution across the FOV afforded by the corrected probes could be useful not only to increase the SNR of neural calcium signals, but also to better isolate individual neurons, decreasing signal contamination across neighboring cells. To test this hypothesis, we considered correlations in “adjacent” simulated neurons, defined as pairs of neurons for which the distance between their centroids was ≤ 25 µm (the radius of the soma of excitatory neurons in the olfactory cortex is around 10 µm 30). Compared to uncorrected probes, we found that corrected microendoscopes of both lengths led to lower fraction of pairs of adjacent cells whose activity correlated significantly more than expected (Figure 6A, F, and Supplementary Table 5), confirming reduced cross contamination among cells. The fraction of pairs of adjacent cells (out of the total number of adjacent pairs) whose activity correlated significantly more than expected increased as a function of the peak SNR threshold for corrected and uncorrected microendoscopes of both lengths (Figure 6A, F). This effect was due to a larger decrease of the total number of pairs of adjacent cells as a function of the peak SNR threshold compared to the decrease in the number of pairs of adjacent cells whose activity was more correlated than expected (Figure 6-figure supplement 1). Moreover, in uncorrected microendoscopes the pairwise correlation of adjacent cells showed a larger linear dependence on the radial distance of the pair from the center of the FOV compared to corrected microendoscopes (Figure 6-figure supplement 2A, C, see also Materials and Methods). We interpreted these findings on correlated neuronal activity as due to the enlarged PSF of uncorrected microendoscopes in lateral portion of the FOV (Figure 3G-J), leading to increased signal contamination by sampling across neurons. These results thus suggest that corrected microendoscopes decrease the artefactual correlation of nearby neurons by diminishing the spatial extension of the excitation volume used for sampling neuronal activity (PSF, Figure 3G-J).

Long corrected microendoscopes enable more precise collection of simulated activity signals from individual cellular sources and decrease cross-contamination between adjacent cells.

Earlier version: A) Fraction of adjacent cell pairs (distance between detected cell centroids ≤ 25 µm) that are more correlated that expected (expected pair correlation was estimated as mean Pearson’s correlation between ground truth activity traces of any possible neuronal pairs plus 3 standard deviations) as a function of the peak SNR threshold imposed on extracted activity traces for n = 13 simulated experiments with the 6.4 mm-long uncorrected (red) and corrected (blue) microendoscope. B) 2D projection of the intersection between the 3D FOV and the 3D ground truth distribution of light sources for a representative uncorrected (left) and corrected (right) synthetic t-series obtained with a 6.4 mm-long microendoscope. The color scale shows the number of overlapping sources that are projected on the same pixel. White boxes mark cell identities detected using CITE-ON 32 without any threshold on the peak SNR of activity traces. Scale bars: 50 μm. Left: the scale applies to the entire FOV. Right: the scale bar refers to the center of the FOV; off-axis scale bar at any radial distance (x and y axes) is locally determined multiplying the length of the scale bar by the corresponding normalized magnification factor shown in the horizontal color-coded bar placed below the image. C) Purity index of extracted traces with peak SNR > 10 was estimated using a GLM of ground truth source contributions and plotted as a function of the radial distance of cell identities from the center of the FOV for n = 13 simulated experiments with the 6.4 mm-long uncorrected (red) and corrected (blue) microendoscope. Black lines represent the linear regression of data ± 95% confidence intervals (shaded colored areas). Slopes ± standard error (s.e.): uncorrected, (-0.0020 ± 0.0002) μm-1; corrected, (-0.0006 ± 0.0001) μm-1. Uncorrected, n = 1365; corrected, n = 1156. Statistical comparison of slopes, p < 10-10, permutation test. D) Distribution of the Pearson’s correlation of extracted activity traces with the first (most correlated) ground truth source for n = 13 simulated experiments with the 6.4 mm-long uncorrected (red) and corrected (blue) microendoscope. Median values: uncorrected, 0.25; corrected, 0.73; the p is computed using the Mann-Whitney U test. E) Pearson’s correlation ± s.e.m. of extracted activity traces with the first (most correlated) ground truth emitter as a function of the radial distance for n = 13 simulated experiments with the 6.4 mm-long uncorrected (red) and corrected (blue) microendoscope. F-J) Same as (A-E) for n = 15 simulated experiments with the 8.8 mm-long uncorrected and corrected microendoscope. H) Slopes ± s.e.: uncorrected, (- 0.0031 ± 0.0003) μm-1; corrected, (-0.0010 ± 0.0002) μm-1. Uncorrected, n = 808; corrected, n = 1328. Statistical comparison of slopes, p < 10-10, permutation test. I) Median values: uncorrected, 0.43; corrected, 0.46; the p is computed using the Mann-Whitney U test.

Revised version: A) Fraction of adjacent cell pairs (distance between detected cell centroids ≤ 25 µm) that are more correlated that expected (expected pair correlation was estimated as mean Pearson’s correlation between ground truth activity traces of any possible neuronal pairs plus 3 standard deviations (SDs)) as a function of the peak SNR threshold imposed on extracted activity traces for n = 13 simulated experiments with the 6.4 mm-long uncorrected (red) and corrected (blue) microendoscope. B) 2D projection of the intersection between the 3D FOV and the 3D ground truth distribution of light sources for a representative uncorrected (left) and corrected (right) synthetic t-series obtained with a 6.4 mm-long microendoscope. The color scale shows the number of overlapping sources that are projected on the same pixel. White boxes mark cell identities detected using CITE-ON 32 without any threshold on the peak SNR of activity traces. Scale bars: 50 μm. Left: the scale applies to the entire FOV. Right: the scale bar refers to the center of the FOV; off-axis scale bar at any radial distance (x and y axes) is locally determined multiplying the length of the scale bar by the corresponding normalized magnification factor shown in the horizontal color-coded bar placed below the image. C) Purity index of extracted traces with peak SNR > 10 was estimated using a GLM of ground truth source contributions and plotted as a function of the radial distance of cell identities from the center of the FOV for n = 13 simulated experiments with the 6.4 mm-long uncorrected (red) and corrected (blue) microendoscope. Black lines represent the linear regression of data ± 95% confidence intervals (shaded colored areas). Slopes ± standard error (s.e.): uncorrected, (−0.0020 ± 0.0002) μm-1; corrected, (−0.0006 ± 0.0001) μm-1. Uncorrected, n = 1365; corrected, n = 1156. Statistical comparison of slopes, p < 10-10, permutation test. Linear regression was repeated using the same range of radial distances for the uncorrected and corrected case, with the maximum value on the x-axis corresponding to the minimum value between the two maximum radial distances obtained in the uncorrected and corrected case (maximum radial distance: 151.6 µm); slopes ± s.e.: uncorrected, (−0.0015 ± 0.0002) µm-1; corrected, (−0.0006 ± 0.0001) μm-1. Uncorrected, n = 991; corrected, n = 1156. Statistical comparison of slopes, p < 10-10, permutation test. D) Distribution of the Pearson’s correlation value of extracted activity traces with the first (most correlated) ground truth source for n = 13 simulated experiments with the 6.4 mm-long uncorrected (red) and corrected (blue) microendoscope. Median values: uncorrected, 0.25; corrected, 0.73; the p is computed using the Mann-Whitney U test. E) Pearson’s correlation ± s.e.m. of extracted activity traces with the first (most correlated) ground truth emitter as a function of the radial distance for n = 13 simulated experiments with the 6.4 mm-long uncorrected (red) and corrected (blue) microendoscope. F-J) Same as (A-E) for n = 15 simulated experiments with the 8.8 mm-long uncorrected and corrected microendoscope. H) Slopes ± s.e.: uncorrected, (−0.0031 ± 0.0003) μm-1; corrected, (−0.0010 ± 0.0002) μm-1. Uncorrected, n = 808; corrected, n = 1328. Statistical comparison of slopes, p < 10-10, permutation test. Linear regression using the same range of radial distances for the uncorrected and corrected case (maximum radial distance: 142.1 μm), slopes ± s.e.: uncorrected, (−0.0014 ± 0.0003) μm-1; corrected, (−0.0010 ± 0.0002) µm-1. Uncorrected, n = 718; corrected, n = 1328. Statistical comparison of slopes, p = 0.0082, permutation test. I) Median values: uncorrected, 0.43; corrected, 0.46; the p is computed using the Mann-Whitney U test.

Earlier version: Using synthetic data, we could precisely quantify the extent to which correcting aberration in microendoscopes decreased mixing signals across neurons. For each ROI, we computed a general linear model (GLM) of the contributions to fluorescence signals of the activity of all ground truth simulated neurons, called “sources”, Fig. 6B, C, G, H). From the GLM coefficients, we extracted a source “purity” index, which reached its highest possible value of 1 when only one source contributed to the calcium fluorescent trace in the considered ROI and had values < 1 when different source neurons were mixed in the fluorescent signal of the considered ROI. For uncorrected microendoscopes of both lengths, we found that purity rapidly decreased with the radial distance of the ROI from the center of the FOV. In contrast, for corrected microendoscopes the purity index decreased far more slowly with the radial distance and the slope of the linear fit of the purity distribution as a function of the radial distance was significantly smaller than that of the uncorrected probe (Fig. 6C, H). For any value of the SNR threshold used to select neurons, the purity index was also larger for corrected than for uncorrected probes (Supplementary Figure 6A, C). Moreover, for each ROI we analyzed the correlation of the extracted traces with the most correlated ground truth neuron (“first ground truth source”) as a function of the radial distance of the ROI (Fig. 6D, E, I, J). For microendoscopes of both lengths, we observed that ROIs from t- series generated with corrected microendoscopes displayed larger correlation with the ground truth signal of an individual neuron over longer distances from the optical axis, compared to uncorrected probes (Fig. 6E, J). Figure 6J showed a moderate enlargement of the distance at which the Pearson’s correlation with the first ground truth emitter started to drop. This may appear at first in contrast with the enlargement of the FOV shown in Figure 4B. However, it can be understood considering that in Figure 4B, all illuminated neurons were visible regardless of whether they were imaged with high axial resolution (e.g., < 10 µm as defined in Figure 3J) or poor axial resolution. In contrast, in Figure 6J we evaluated the correlation between the calcium trace extracted from a given ROI and the real activity trace of the first simulated ground truth emitter for that specific ROI. The moderate increase in the correlation for the corrected microendoscope compared to the uncorrected microendoscope (Figure 6J) was consistent with the moderate improvement in the axial resolution of the corrected probe compared to the uncorrected probe at intermediate radial distances (60-100 µm from the optical axis, see Figure 3J). Taken together, these results indicated that corrected microendoscopes enabled more precise recording of individual cellular source neurons and enhanced demixing of signals coming from different neurons.

Revised version: Using synthetic data, we could precisely quantify the extent to which correcting aberration in microendoscopes decreased mixing signals across neurons. For each ROI, we computed a general linear model (GLM) of the contributions to fluorescence signals of the activity of all ground truth simulated neurons, called “sources”, Figure 6B, C, G, H). From the GLM coefficients, we extracted a source “purity” index, which reached its highest possible value of 1 when only one source contributed to the calcium fluorescent trace in the considered ROI and had values < 1 when different source neurons were mixed in the fluorescent signal of the considered ROI. For uncorrected microendoscopes of both lengths, we found that purity rapidly decreased with the radial distance of the ROI from the center of the FOV. In contrast, for corrected microendoscopes the purity index decreased far more slowly with the radial distance and the slope of the linear fit of the purity distribution as a function of the radial distance was significantly smaller than that of the uncorrected probe (Figure 6C, H). For any value of the peak SNR threshold used to select neurons, the purity index was also larger for corrected than for uncorrected probes (Figure 6-figure supplement 2B, D). Moreover, for each ROI we analyzed the correlation of the extracted traces with the most correlated ground truth neuron (“first ground truth source”) as a function of the radial distance of the ROI (Figure 6D, E, I, J). For microendoscopes of both lengths, we observed that ROIs from t-series generated with corrected microendoscopes displayed larger correlation with the ground truth signal of an individual neuron over longer distances from the optical axis, compared to uncorrected probes (Figure 6E, J). Figure 6J showed a moderate enlargement of the distance at which the Pearson’s correlation with the first ground truth emitter started to drop. This may appear at first in contrast with the enlargement of the FOV shown in Figure 4B. However, it can be understood considering that in Figure 4B, all illuminated neurons were visible regardless of whether they were imaged with high axial resolution (e.g., < 10 µm as defined in Figure 3J) or poor axial resolution. In contrast, in Figure 6J we evaluated the correlation between the calcium trace extracted from a given ROI and the real activity trace of the first simulated ground truth emitter for that specific ROI. The moderate increase in the correlation for the corrected microendoscope compared to the uncorrected microendoscope (Figure 6J) was consistent with the moderate improvement in the axial resolution of the corrected probe compared to the uncorrected probe at intermediate radial distances (60-100 µm from the optical axis, see Figure 3J). Taken together, these results indicated that corrected microendoscopes enabled more precise recording of individual cellular source neurons and enhanced demixing of signals coming from different neurons.

Validation of long corrected microendoscopes for large FOV population imaging in ventral brain region of awake mice

Earlier version: We tested long corrected microendoscopes for in vivo 2P calcium imaging experiments in ventral regions of the mouse brain. To this aim, we implanted mice with corrected microendoscopes based on 8.8 mm-long GRIN lenses over the mouse anterior olfactory cortex, a deep region laying > 4 mm down the mouse brain surface (Fig. 7A, B). Mice expressed genetically encoded calcium indicators, either jGCaMP7f or jGCaMP8f, in excitatory neurons through adeno-associated virus (AAVs) injection. In awake head-fixed mice, we recorded 2P imaging t-series (λexc = 920 nm; frame rate: 30 Hz; t-series duration: 8 minutes) during repetitive olfactory stimulation with various odorants (Fig. 7C-E, see Materials and Methods).

Revised version: We tested long corrected microendoscopes for in vivo 2P calcium imaging experiments in ventral regions of the mouse brain. To this aim, we implanted mice with corrected microendoscopes based on 8.8 mm-long GRIN lenses over the mouse anterior olfactory cortex, a deep region laying > 4 mm down the mouse brain surface (Figure 7A, B). Mice expressed genetically encoded calcium indicators, either jGCaMP7f or jGCaMP8f, in excitatory neurons through adeno-associated virus (AAVs) injection. In awake head-fixed mice, we recorded 2P imaging t-series (λexc = 920 nm; frame rate: 30 Hz; t-series duration: 8 minutes) during repetitive olfactory stimulation with various odorants (Figure 7C-E, see Materials and Methods).

Enlarged FOV population imaging of ventral regions of the brain with long corrected microendoscopes in awake mice.

Earlier version: A) Schematic showing injection of the viral solution in the mouse piriform cortex. B) Representative section of fixed brain tissue showing the position of the GRIN lens implant. jGCaMP7f fluorescence is shown in red. NeuroTrace (Neurotrace) Nissl staining is shown in cyan. C) Schematic showing the experimental preparation for 2P corrected microendoscope imaging in awake mice. D) Top and bottom: example of FOV with excitatory neurons expressing the calcium sensor jGCaMP7f, obtained using the 8.8 mm-long corrected microendoscope in the mouse PC. The bottom image shows the rectangular boxes (green) indicating the position of detected neurons generated by CITE-ON (see Materials and Methods, 32). On-axis scale bar: 20 µm; for off-axis scale bar, refer to the magnification factor bar below the image (see Materials and Methods). E) Representative jGCaMP7f activity traces extracted from the recording shown in (D) through eight minute-long continuous 2P imaging recordings (λexc = 920 nm) in an awake head-fixed mouse. Vertical dashed orange lines mark the onset and the end of the 22 s-long trials in which the imaging session was divided. F) Examples of neuronal responses to different olfactory stimuli for three representative cells. For each odor (top labels), gray lines represent calcium responses recorded in single trials and the red line represent the average response. Light blue areas indicate the time interval of stimulus presentation.

Revised version: A) Schematic showing injection of the viral solution in the mouse piriform cortex. B) Representative section of fixed brain tissue showing the position of the GRIN lens implant. jGCaMP7f fluorescence is shown in red. NeuroTrace (Neurotrace) Nissl staining is shown in cyan. C) Schematic showing the experimental preparation for 2P corrected microendoscope imaging in awake mice. D) Top and bottom: example of FOV with excitatory neurons expressing the calcium sensor jGCaMP7f, obtained using the 8.8 mm-long corrected microendoscope in the mouse piriform cortex. The bottom image shows the rectangular boxes (green) indicating the position of detected neurons generated by CITE-ON (see Materials and Methods, 32). On-axis scale bar: 20 µm; for off-axis scale bar, refer to the magnification factor bar below the image (see Materials and Methods). E) Representative jGCaMP7f activity traces extracted from the recording shown in (D) through eight minute-long continuous 2P imaging recordings (λexc = 920 nm) in an awake head-fixed mouse. Vertical dashed orange lines mark the onset and the end of the 22 s-long trials in which the imaging session was divided. F) Examples of neuronal responses to different olfactory stimuli for three representative cells. For each odor (top labels), gray lines represent calcium responses recorded in single trials and the red line represent the average response. Light blue areas indicate the time interval of stimulus presentation.

Earlier version: In recorded t-series, we first observed that the SNR of calcium events recorded in individual cells was constant across the FOV (Fig. 8A). We then computed the pairwise correlation of the activity traces of pairs of adjacent neurons (defined as pairs of cells with an intersomatic distance of < 25 µm) as a function of the radial distance of the pair’s position centroid from the center of the FOV (Fig. 8B). We found no dependence of the pairwise correlation on the distance from the center of the FOV. In fact, the slope of the linear fit of pairwise correlation as a function of the radial distance was not significantly different from zero (Fig. 8B). These findings were in agreement with the results of the optical characterization (Fig. 3) and with those of the simulation of calcium data (Fig. 5, 6). These results suggested that the absence of dependence of the pairwise correlation with radial distance reflected the improved and more homogeneous spatial resolution of corrected microendoscopes, which prevented source mixing and artifically inflated correlations of adjacent neurons in the FOV periphery. Varying the distance between centroids of adjacent neurons did not change the result (Supplementary Table 6). Running the analysis separately for mice expressing jGCaMP8f and jGCaMP7f also confirmed the previous findings (Supplementary Table 6). Finally, the Pearson’s correlation coefficient between the activity traces of any pair of neurons within the FOV decresed as a function of their reciprocal distance (Fig. 8C). This decrease in correlation could reflect a slight change in the laminar position of recorded neurons or suggest a decrease in recurrent connectivity between neurons of the olfactory cortex with distance 33, 34.

Revised version: In recorded t-series, we first observed that the peak SNR of calcium events recorded in individual cells was constant across the FOV (Figure 8A). We then computed the pairwise correlation of the activity traces of pairs of adjacent neurons (defined as pairs of cells with an intersomatic distance of ≤ 25 µm) as a function of the radial distance of the pair’s position centroid from the center of the FOV (Figure 8B). We found no dependence of the pairwise correlation on the distance from the center of the FOV. In fact, the slope of the linear fit of pairwise correlation as a function of the radial distance was not significantly different from zero (Figure 8B). These findings were in agreement with the results of the optical characterization (Figure 3) and with those of the simulation of calcium data (Figure 5, 6). These results suggested that the absence of dependence of the pairwise correlation with radial distance reflected the improved and more homogeneous spatial resolution of corrected microendoscopes, which prevented source mixing and artifically inflated correlations of adjacent neurons in the FOV periphery. Varying the distance between centroids of adjacent neurons did not change the result (Supplementary Table 6). Running the analysis separately for mice expressing jGCaMP8f and jGCaMP7f also confirmed the previous findings (Supplementary Table 6). Finally, the Pearson’s correlation coefficient between the activity traces of any pair of neurons within the FOV decresed as a function of their reciprocal distance (Figure 8C). This decrease in correlation could reflect a slight change in the laminar position of recorded neurons or suggest a decrease in recurrent connectivity between neurons of the olfactory cortex with distance 33, 34.

Unbiased population imaging in the piriform cortex of awake mice using long corrected microendoscopes.

Earlier version: A) Left: schematic representation of trace parameters used to compute peak SNR. The red solid line marks the maximum value of the trace, the blue line indicates the standard deviation of intensity values below the 25th percentile of the intensity distribution of the entire trace (see Materials and Methods), and the dashed grey line marks the origin of the y-axis (ΔF/F0 = 0). Right: peak SNR of calcium traces extracted from individual cells as a function of radial distance of the cell (n = 240 neurons from m = 6 FOV). The red line is the linear regression of data points: intercept ± s.e. = 19 ± 2; slope ± s.e. = (0.03 ± 0.02) μm-1. The slope is not significantly different from zero, p = 0.18, permutation test. B) Left: schematic representation showing pairs of neurons and distance definitions for the Pearson’s correlation analysis shown on the right (in red, distance between “adjacent neurons” dpair ≤ 25 µm). Right: Pearson’s correlation of calcium traces from “adjacent neurons” as a function of radial distance of the pair centroid from the center of the FOV (n = 195 adjacent neuron pairs from m = 6 FOV). The red line represents the linear regression of data points: intercept ± s.e. = 0.41 ± 0.03; slope ± s.e. = (-0.0006 ± 0.0004) μm-1. The slope is not significantly different from zero, p = 0.089, Wald test. C) Left: schematic representation of neuron pairs used for the analysis on the right (any possible dpair). Right: Pearson’s correlation of calcium traces from pairs of neurons as a function of the distance between them (n = 4767 pairs from m = 6 FOV). The red line is the linear regression of data points: intercept ± s.e. = 0.288 ± 0.005; slope ± s.e. = (-87 ± 4) · 10-5 μm-1. The slope is significantly different from zero, p = 2 · 10-5, permutation test.

Revised version: A) Left: schematic representation of trace parameters used to compute peak SNR. The red solid line marks the maximum value of the trace, the blue line indicates the SD of intensity values below the 25th percentile of the intensity distribution of the entire trace (see Materials and Methods), and the dashed grey line marks the origin of the y-axis (ΔF/F0 = 0). Right: peak SNR of calcium traces extracted from individual cells as a function of radial distance of the cell (n = 240 neurons from m = 6 FOVs). The red line is the linear regression of data points: intercept ± s.e. = 19 ± 2; slope ± s.e. = (0.03 ± 0.02) μm-1. The slope is not significantly different from zero, p = 0.18, permutation test. B) Left: schematic representation showing pairs of neurons and distance definitions for the Pearson’s correlation analysis shown on the right (in red, distance between “adjacent neurons” dpair

Combined with the results of the simulated calcium data, these in vivo findings demonstrate that, by decreasing off-axis aberrations of commercial GRIN lenses, long corrected microendoscopes enable unbiased sampling of neural population activity on a large FOV in ventral regions of the mouse brain in vivo.

Discussion

Combining GRIN lens-based endoscopy 8, 9 with 2P fluorescence imaging 3, 35 is a powerful approach to study the physiology of deep brain circuits. However, GRIN lenses are not ideal optical elements and they have intrinsic optical aberrations 18, 19, which significantly degrade the PSF and reduce the effective FOV 7, 20, 21, 23. In the current study, we developed, characterized, and validated new aberration corrected microendoscopic probes to address this critical need. Corrected microendoscopes had length of 6.4 mm and 8.8 mm (see Table 2 and Materials and Methods), allowing reaching the deepest structures of the mouse brain, and cross section of 500 µm, ensuring limited tissue invasiveness.

Earlier version: We first designed the corrected microendoscopes using optical ray-race simulations. Simulations showed strongly aberrated PSF in lateral portions of the FOV of uncorrected microendoscopes and predicted that single corrective optical elements with the aspherical profiles shown in Fig. 1A-B improved the optical performance of microendoscopes of either length. This prediction was experimentally confirmed using subresolved fluorescent beads and layers. We first observed that uncorrected microendoscopes of either length displayed strong astigmatism and had largely distorted PSF in lateral regions of the FOV. In corrected microendoscopes, we instead found that the PSF had regular shape and constrained values for larger radial distances, leading to more homogeneous spatial resolution across the FOV and enlarged effective FOV. Importantly, experimental values describing the improvement of the optical performance of the corrected microendoscopes (Fig. 3G-J and Supplementary Table 4) were generally in good agreement with the prediction of the optical simulations (Fig. 1D, F and Supplementary Table 2). The size of simulated PSFs at a given radial distance (e.g., 90 µm, Figure 1) tended to be generally smaller than that of the experimentally measured PSFs (Figure 3). This might be due to multiple reasons. First, simulated PSFs were excitation PSFs, i.e. they described the intensity spatial distribution of focused excitation light. On the contrary, measured PSFs resulted from the excitation and emission processes and thus they were also affected by aberrations of light emitted by fluorescent beads and collected by the microscope. Second, in the optical simulations first-order and not higher-order aberrations were considered. Third, intrinsic variability in experimental parameters (e.g., alignment of the microendoscope to the optical axis of the microscope, distance between the GRIN back end and the objective) were not considered in the simulations.

Revised version: We first designed the corrected microendoscopes using optical ray-race simulations. Simulations showed strongly aberrated PSF in lateral portions of the FOV of uncorrected microendoscopes and predicted that single corrective optical elements with the aspherical profiles shown in Figure 1A, B improved the optical performance of microendoscopes of either length. This prediction was experimentally confirmed using subresolved fluorescent beads and layers. We first observed that uncorrected microendoscopes of either length displayed strong astigmatism and had largely distorted PSF in lateral regions of the FOV. In corrected microendoscopes, we instead found that the PSF had regular shape and constrained values for larger radial distances, leading to more homogeneous spatial resolution across the FOV and enlarged effective FOV. Importantly, experimental values describing the improvement of the optical performance of the corrected microendoscopes (Figure 3G-J and Supplementary Table 4) were generally in good agreement with the prediction of the optical simulations (Figure 1D, F and Supplementary Table 2). The size of simulated PSFs at a given radial distance (e.g., 90 µm, Figure 1) tended to be generally smaller than that of the experimentally measured PSFs (Figure 3). This might be due to multiple reasons. First, simulated PSFs were excitation PSFs, i.e. they described the intensity spatial distribution of focused excitation light. On the contrary, measured PSFs resulted from the excitation and emission processes and thus they were also affected by aberrations of light emitted by fluorescent beads and collected by the microscope. Second, in the optical simulations first-order and not higher-order aberrations were considered. Third, intrinsic variability in experimental parameters (e.g., alignment of the microendoscope to the optical axis of the microscope, distance between the GRIN back end and the objective) were not considered in the simulations. It is important to note that the improvement in spatial resolution of GRIN lens-based probes described in this work was related exclusively to the introduction of the corrective lens. In fact, uncorrected microendoscopes used for comparison consisted of a GRIN rod and a glass coverslip attached to the back end of the GRIN lens. Previous work showed that the addition of a glass coverslip, that typically introduces negative spherical aberrations, counteracted positive spherical aberrations introduced by GRIN lens singlets and improved the optical performance of the GRIN rod when combined with the use of a microscope objective with corrective collar 36. Therefore, it is reasonable to speculate that adopting our corrected microendoscopes rather than using bare GRIN lenses enhances imaging performance even more than what described in the present study.

Earlier version: Did the improved optical properties of corrected microendoscopes increase the quality of imaging when corrected microendoscopes were applied to biological samples? We directly addressed this question by performing 2P imaging with uncorrected and corrected microendoscopes in brain slices in which neurons expressed a green fluorescent indicator. In uncorrected microendoscopes, we found that neuronal cell bodies were bright and well contrasted only in the center of the FOV and that neuronal cell bodies were dimmer and blurred in lateral regions of the FOV (Fig. 4). Moreover, neuronal processes could be observed only in the center of the FOV and not in lateral portions of the FOV (Fig. 4). In contrast, in corrected microendoscopes of either length, we found that neurons were bright and well contrasted across the whole FOV, which enabled distinguishing neurons, which were blurred and merged together when imaged with uncorrected microendoscopes (Fig. 4). Furthermore, small neuronal processes could be observed throughout the entire FOV. These imaging findings in brain slices are in line with the results of the optical simulation and of the experimental characterization, which showed improved PSF and more homogeneous spatial resolution in corrected microendoscopes compared to uncorrected ones. It must be considered, however, that the extended FOV achieved by our aberration correction method was characterized by a curved focal plane. Therefore, cells located in different radial positions within the image were located at different axial positions and cells at the border of the FOV were closer to the front end of the microendoscope.

Revised version: Did the improved optical properties of corrected microendoscopes increase the quality of imaging when corrected microendoscopes were applied to biological samples? We directly addressed this question by performing 2P imaging with uncorrected and corrected microendoscopes in brain slices in which neurons expressed a green fluorescent indicator. In uncorrected microendoscopes, we found that neuronal cell bodies were bright and well contrasted only in the center of the FOV and that neuronal cell bodies were dimmer and blurred in lateral regions of the FOV (Figure 4). Moreover, neuronal processes could be observed only in the center of the FOV and not in lateral portions of the FOV (Figure 4). In contrast, in corrected microendoscopes of either length, we found that neurons were bright and well contrasted across the whole FOV, which enabled distinguishing neurons, which were blurred and merged together when imaged with uncorrected microendoscopes (Figure 4). Furthermore, small neuronal processes could be observed throughout the entire FOV. These imaging findings in brain slices are in line with the results of the optical simulation and of the experimental characterization, which showed improved PSF and more homogeneous spatial resolution in corrected microendoscopes compared to uncorrected ones. It must be considered, however, that the extended FOV achieved by our aberration correction method was characterized by a curved focal plane. Therefore, cells located in different radial positions within the image were located at different axial positions and cells at the border of the FOV were closer to the front end of the microendoscope.

Earlier version: How does the improved optical performance of corrected microendoscope translate into enhanced ability to extract meaningful information from population of neurons in functional fluorescence imaging experiments? We addressed this question generating and analyzing synthetic calcium data in which we could compare the signals extracted from simulated recordings using uncorrected and corrected microendoscopes with the ground truth signals, which were used to generate the simulations. We found that the aberration correction led to an increase in the number of ROI with high SNR and, correspondingly, a decrease in the number of ROIs with low SNR. Moreover, aberration correction led to high and more stable purity of source demixing over radial distance, which, in turn, led to a major decrease in the artificial overestimation of correlations between adjacent neurons due to source mixing. Since investigating neuronal population codes needs quantification of neuronal activity with large SNR, high precision, and limited contamination, we expect the corrected microendoscopes presented in this manuscript will greatly contribute to the investigation of deep brain regions. In particular, the reduction of artificial overestimation of correlations due to source mixing appears crucial to precisely evaluate within-network neuron-to-neuron communication and investigate the role of correlations between neurons in population codes. Artificial correlations between nearby neurons would bias estimates of neuron-to-neuron communication toward overestimating communications at short distances. Moreover, because source mixing duplicates one neuron into another and thus can only increase redundancy, artificial correlations would also bias the effect of correlations on population coding by artificially inflating redundant correlations over synergistic interactions between neurons 36, 37. Thus, correcting optical aberrations in GRIN lenses as presented in this study appears essential for the unbiased characterization of single-neuron and population codes in deeper brain structures.

Revised version: How does the improved optical performance of corrected microendoscope translate into enhanced ability to extract meaningful information from population of neurons in functional fluorescence imaging experiments? We addressed this question generating and analyzing synthetic calcium data in which we could compare the signals extracted from simulated recordings using uncorrected and corrected microendoscopes with the ground truth signals, which were used to generate the simulations. We found that the aberration correction led to an increase in the number of ROIs with high SNR and, correspondingly, a decrease in the number of ROIs with low SNR. Moreover, aberration correction led to high and more stable purity of source demixing over radial distance, which, in turn, led to a major decrease in the artificial overestimation of correlations between adjacent neurons due to source mixing. Since investigating neuronal population codes needs quantification of neuronal activity with large SNR, high precision, and limited contamination, we expect the corrected microendoscopes presented in this manuscript will greatly contribute to the investigation of deep brain regions. In particular, the reduction of artificial overestimation of correlations due to source mixing appears crucial to precisely evaluate within-network neuron-to-neuron communication and investigate the role of correlations between neurons in population codes. Artificial correlations between nearby neurons would bias estimates of neuron-to-neuron communication toward overestimating communications at short distances. Moreover, because source mixing duplicates one neuron into another and thus can only increase redundancy, artificial correlations would also bias the effect of correlations on population coding by artificially inflating redundant correlations over synergistic interactions between neurons 37, 38. Thus, correcting optical aberrations in GRIN lenses as presented in this study appears essential for the unbiased characterization of single-neuron and population codes in deeper brain structures.

Earlier version: We used long corrected microendoscopes to measure population dynamics in the olfactory cortex of awake head-restrained mice with unprecedented combination of high spatial resolution across the FOV and minimal invasiveness 17. In agreement with the results of our optical characterization and of the synthetic calcium data, we found that the SNR of calcium signals was independent on the radial distance and that high SNR calcium signals could be recorded across the whole FOV. Moreover, we observed that the correlation between calcium signals of adjacent neurons did not depend on the radial distance, in line with the more homogeneous PSF that characterizes corrected microendoscopes. The correlation between the calcium signals of any pair of neurons displayed, instead, a negative dependence on the distance between the two neurons. This decrease in correlation could reflect a slight change in the laminar position of recorded neurons or suggest a decrease in recurrent connectivity between neurons of the olfactory cortex with distance 33, 34. In relationship to studies on the olfactory cortex, the improved optical characteristics of corrected microendoscopes increase our ability to understand the properties of this network in several fundamental ways. First, imaging more cells with increased SNR improves our ability to decipher information encoded by neural ensembles in the olfactory cortex. For example, we will be able to characterize the odor response properties of genetically defined subpopulations of olfactory cortex neurons, which may be too sparse to be reliably captured in the smaller effective FOV of uncorrected GRIN lenses. Second, imaging more cells with high SNR in each mouse will allow us to explore differences between individual mice, rather than pooling cells from different animals. Third, less cross-contamination of fluorescence signals improves cell segmentation and motion correction, which are major challenges when imaging mice that are awake or actively engaged in behavioral tasks 38, 39. Finally, improved cell segmentation facilitates cell tracking in chronic imaging experiments. Thus, corrected microendoscopes will enhance our ability to monitor individual cells over extended periods of time, for example while mice learn to perform an odor discrimination task 40.

Revised version: We used long corrected microendoscopes to measure population dynamics in the olfactory cortex of awake head-restrained mice with unprecedented combination of high spatial resolution across the FOV and minimal invasiveness 17. In agreement with the results of our optical characterization and of the synthetic calcium data, we found that the SNR of calcium signals was independent on the radial distance and that high SNR calcium signals could be recorded across the whole FOV. Moreover, we observed that the correlation between calcium signals of adjacent neurons did not depend on the radial distance, in line with the more homogeneous PSF that characterizes corrected microendoscopes. The correlation between the calcium signals of any pair of neurons displayed, instead, a negative dependence on the distance between the two neurons. This decrease in correlation could reflect a slight change in the laminar position of recorded neurons or suggest a decrease in recurrent connectivity between neurons of the olfactory cortex with distance 33, 34. In relationship to studies on the olfactory cortex, the improved optical characteristics of corrected microendoscopes increase our ability to understand the properties of this network in several fundamental ways. First, imaging more cells with increased SNR improves our ability to decipher information encoded by neural ensembles in the olfactory cortex. For example, we will be able to characterize the odor response properties of genetically defined subpopulations of olfactory cortex neurons, which may be too sparse to be reliably captured in the smaller effective FOV of uncorrected GRIN lenses. Second, imaging more cells with high SNR in each mouse will allow us to explore differences between individual mice, rather than pooling cells from different animals. Third, less cross-contamination of fluorescence signals improves cell segmentation and motion correction, which are major challenges when imaging mice that are awake or actively engaged in behavioral tasks 39, 40. Finally, improved cell segmentation facilitates cell tracking in chronic imaging experiments. Thus, corrected microendoscopes will enhance our ability to monitor individual cells over extended periods of time, for example while mice learn to perform an odor discrimination task 41.

Earlier version: Alternative methods to correct optical aberrations in GRIN lenses include, for example, low-order adaptive optics 18, adaptive optics through pupil segmentation 20, 21, and geometric transformation adaptive optics 23. In the first approach, an adaptive optics module comprising an electrostatic membrane mirror was introduced in the optical path of the 2P microscope and it was shown to significantly improve the microscope PSF 18. In adaptive optics through pupil segmentation instead, the objective rear lens was divided into N different subsections. By comparing images acquired with light going through each individual pupil segment with images acquired with light going through the whole pupil, local wavefront tilts introduced by aberrations were calculated and corrected for using a spatial light modulator (SLM) 41. This method was applied to GRIN lens-based endoscopes of larger diameter (i.e., 1.4 mm) compared to the one used in this study (i.e., 0.5 mm) and it was shown to enlarge the effective FOV during 2P imaging 20, 21. Finally, in geometric transformation adaptive optics, a 90-degrees rotation was applied to the aberrated beam spatial profile exiting a first GRIN lens and the aberrated beam was passed through a second GRIN lens. This strategy was effective in correcting chromatic and volumetric aberrations in GRIN microendoscopes of small cross section (500 µm 23). However, all adaptive optics methods described above require substantial change the optical set-up (e.g., the insertion of an SLM, membrane mirror, and additional optics) and the development of specific software control 18, 20, 21,23. Additionally, under certain conditions pupil segmentation may limit the acquisition frame rate of the system 21. In contrast to the methods described above, the long corrected microendoscopes presented in the current manuscript do not limit the temporal resolution of the acquisition system, require no modification of the hardware and software of the optical set-up, and are easily coupled to any standard 2P microscope. Another advantage of long corrected microendoscopes described here over adaptive optics approaches is the possibility to couple corrected microendoscopes with portable 2P microscopes 42–44, allowing high resolution functional imaging of deep brain circuits on an enlarged FOV during naturalistic behavior in freely moving mice. The corrected microendoscopes presented in this study thus represent a ready-to-use solution, which will likely ease the adoption of aberration correction in GRIN lens-based 2P functional imaging across experimental laboratories.

Revised version: Alternative methods to correct optical aberrations in GRIN lenses include, for example, low-order adaptive optics 18, adaptive optics through pupil segmentation 20, 21, and geometric transformation adaptive optics 23. In the first approach, an adaptive optics module comprising an electrostatic membrane mirror was introduced in the optical path of the 2P microscope and it was shown to significantly improve the microscope PSF 18. In adaptive optics through pupil segmentation instead, the objective rear lens was divided into N different subsections. By comparing images acquired with light going through each individual pupil segment with images acquired with light going through the whole pupil, local wavefront tilts introduced by aberrations were calculated and corrected for using a spatial light modulator (SLM) 42. This method was applied to GRIN lens-based endoscopes of larger diameter (i.e., 1.4 mm) compared to the one used in this study (i.e., 0.5 mm) and it was shown to enlarge the effective FOV during 2P imaging 20, 21. Finally, in geometric transformation adaptive optics, a 90-degrees rotation was applied to the aberrated beam spatial profile exiting a first GRIN lens and the aberrated beam was passed through a second GRIN lens. This strategy was effective in correcting chromatic and volumetric aberrations in GRIN microendoscopes of small cross section (500 µm 23). However, all adaptive optics methods described above require substantial change the optical set-up (e.g., the insertion of an SLM, membrane mirror, and additional optics) and the development of specific software control 18, 20, 21, 23. Additionally, under certain conditions pupil segmentation may limit the acquisition frame rate of the system 21. In contrast to the methods described above, the long corrected microendoscopes presented in the current manuscript do not limit the temporal resolution of the acquisition system, require no modification of the hardware and software of the optical set-up, and are easily coupled to any standard 2P microscope. Another advantage of long corrected microendoscopes described here over adaptive optics approaches is the possibility to couple corrected microendoscopes with portable 2P microscopes 43–45, allowing high resolution functional imaging of deep brain circuits on an enlarged FOV during naturalistic behavior in freely moving mice. The corrected microendoscopes presented in this study thus represent a ready-to-use solution, which will likely ease the adoption of aberration correction in GRIN lens-based 2P functional imaging across experimental laboratories.

Earlier version: In conclusion, we designed, developed, and validated a new series of aberration-corrected GRIN lens-based microendoscopes, which are long enough to reach the most ventral regions of the mouse brain. Long corrected microendoscopes had improved axial resolution and several folds enlarged effective FOV. Importantly, these optical improvements enabled more precise extraction of single cell and population signals from 2P imaging recordings. We validated long corrected microendoscopes by performing functional imaging of the olfactory cortex circuits in awake head-fixed mice. Long corrected microendoscope provide a unique combination of high and homogenous optical properties across the FOV and minimal invasiveness. Furthermore, they do not require modification of the hardware and software of the optical microscope, providing a ready-to-use solution to experimental laboratories. We foresee that the long corrected microendoscopes developed in this study will drastically improve the yield of 2P imaging in deep brain areas of mice and larger mammals, such as rats 45, cats 46, 47, and marmosets 48, 49.

Revised version: In conclusion, we designed, developed, and validated a new series of aberration-corrected GRIN lens-based microendoscopes, which are long enough to reach the most ventral regions of the mouse brain. Long corrected microendoscopes had improved axial resolution and several folds enlarged effective FOV. Importantly, these optical improvements enabled more precise extraction of single cell and population signals from 2P imaging recordings. We validated long corrected microendoscopes by performing functional imaging of the olfactory cortex circuits in awake head-fixed mice. Long corrected microendoscopes provide a unique combination of high and homogenous optical properties across the FOV and minimal invasiveness. Furthermore, they do not require modification of the hardware and software of the optical microscope, providing a ready-to-use solution to experimental laboratories. We foresee that the long corrected microendoscopes developed in this study will drastically improve the yield of 2P imaging in deep brain areas of mice and larger mammals, such as rats 46, cats 47, 48, and marmosets 49, 50.